Tumors mimic the niche to inhibit neighboring stem cell differentiation

Author response:

The following is the authors’ response to the previous reviews

eLife Assessment

This study presents results supporting a model that tumorous germline stem cells (GSCs) in the Drosophila ovary mimic the stem cell niche and inhibit the differentiation of neighboring cells. The valuable findings show that GSC tumors often contain non-mutant cells whose differentiation is suppressed by the GSC tumorous cells. However, the evidence showing that the GSC tumors produce BMP ligands to suppress differentiation of non-mutant cells is incomplete due to concerns about the new HCR data.

Thanks for this assessment. All concerns raised by the reviewers regarding the HCR data and others are followed by our responses below.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This preprint from Shaowei Zhao and colleagues presents results that suggest tumorous germline stem cells (GSCs) in the Drosophila ovary mimic the ovarian stem cell niche and inhibit the differentiation of neighboring non-mutant GSC-like cells. The authors use FRT-mediated clonal analysis driven by a germline-specific gene (nos-Gal4, UASp-flp) to induce GSC-like cells mutant for bam or bam's co-factor bgcn. Bam-mutant or bgcn-mutant germ cells produce tumors in the stem cell compartment (the germarium) of the ovary (Fig. 1). These tumors contain non-mutant cells - termed SGC for single-germ cells. 75% of SGCs do not exhibit signs of differentiation (as assessed by bamP-GFP) (Fig. 2). The authors demonstrate that block in differentiation in SGC is a result of suppression of bam expression (Fig. 2). They present data suggesting that in 73% of SGCs BMP signaling is low (assessed by dad-lacZ) (Fig. 3) and proliferation is less in SGCs vs GSCs. They present genetic evidence that mutations in BMP pathway receptors and transcription factors suppress some of the non-autonomous effects exhibited by SGCs within bam-mutant tumors (Fig. 4). They show data that bam-mutant cells secrete Dpp, but this data is not compelling (see below) (Fig. 5). They provide genetic data that loss of BMP ligands (dpp and gbb) suppresses the appearance of SGCs in bam-mutant tumors (Fig. 6). Taken together, their data support a model in which bam-mutant GSC-like cells produce BMPs that act on non-mutant cells (i.e., SGCs) to prevent their differentiation, similar to what in seen in the ovarian stem cell niche. This preprint from Shaowei Zhao and colleagues presents results that suggest tumorous germline stem cells (GSCs) in the Drosophila ovary mimic the ovarian stem cell niche and inhibit the differentiation of neighboring non-mutant GSC-like cells. The authors use FRT-mediated clonal analysis driven by a germline-specific gene (nos-Gal4, UASp-flp) to induce GSC-like cells mutant for bam or bam's co-factor bgcn. Bam-mutant or bgcn-mutant germ cells produce tumors in the stem cell compartment (the germarium) of the ovary (Fig. 1). These tumors contain non-mutant cells - termed SGC for single-germ cells. 75% of SGCs do not exhibit signs of differentiation (as assessed by bamP-GFP) (Fig. 2). The authors demonstrate that block in differentiation in SGC is a result of suppression of bam expression (Fig. 2). They present data suggesting that in 73% of SGCs BMP signaling is low (assessed by dad-lacZ) (Fig. 3) and proliferation is less in SGCs vs GSCs. They present genetic evidence that mutations in BMP pathway receptors and transcription factors suppress some of the non-autonomous effects exhibited by SGCs within bam-mutant tumors (Fig. 4). They show data that bam-mutant cells secrete Dpp, but this data is not compelling (see below) (Fig. 5). They provide genetic data that loss of BMP ligands (dpp and gbb) suppresses the appearance of SGCs in bam-mutant tumors (Fig. 6). Taken together, their data support a model in which bam-mutant GSC-like cells produce BMPs that act on non-mutant cells (i.e., SGCs) to prevent their differentiation, similar to what in seen in the ovarian stem cell niche.

Strengths:

(1) Use of an excellent and established model for tumorous cells in a stem cell microenvironment

(2) Powerful genetics allow them to test various factors in the tumorous vs non-tumorous cells

(3) Appropriate use of quantification and statistics

Thank you for your valuable comments, and we greatly appreciate them.

Weaknesses:

(1) What is the frequency of SGCs in nos>flp; bam-mutant tumors? For example, are they seen in every germarium, or in some germaria, etc or in a few germaria.

This concern was addressed in the rebuttal. The line number is 106, not line 103.

(2) Does the breakdown in clonality vary when they induce hs-flp clones in adults as opposed to in larvae/pupae?

This concern was addressed in the rebuttal. However, these statements are no on lines 331-335 but instead starting on line 339. Please be accurate about the line numbers cited in the rebuttal. They need to match the line numbers in the revised manuscript.

We have rechecked the line numbers and confirmed that the mismatch arose from the Word-to-PDF conversion process on the eLife website. As this issue has recurred and reviewers’ file-format preferences are unknown to us, we have added a clarifying note at the beginning of each response letter: “Please note that the line numbers cited refer to the revised manuscript in the Microsoft Word format”.

(3) Approximately 20-25% of SGCs are bam+, dad-LacZ+. Firstly, how do the authors explain this? Secondly, of the 70-75% of SGCs that have no/low BMP signaling, the authors should perform additional characterization using markers that are expressed in GSCs (i.e., Sex lethal and nanos).

The authors did not perform additional staining for GSC-enriched protein like Sex lethal and nanos.

The 70-75% of SGCs that have low BMP signaling display the following characteristics: 1) dot-like spectrosomes, 2) positivity for Dad-lacZ, and 3) absence of bamP-GFP expression. This combination of traits is sufficient to classify them as GSC-like cells. Neither Sex lethal nor Nanos is expressed exclusively in GSCs (Chau et al., 2009; Li et al., 2009), rendering them unsuitable for distinguishing GSC-like from cystoblast-like cells.

(4) All experiments except Fig. 1I (where a single germarium with no quantification) were performed with nos-Gal4, UASp-flp. Have the authors performed any of the phenotypic characterizations (i.e., figures other than figure 1) with hs-flp?

In the rebuttal, the authors stated that they used nos>flp for all figures except for Fig. 1I. It would be more convincing for them to prove in Fig. 1 than there is not phenoytpic difference between the two methods and then switch to the nos>FLP method for the rest of the paper.

We appreciate this suggestion. These data are included in Figure 1-figure supplement 3 in the revised manuscript.

(5) Does the number of SGCs change with the age of the female? The experiments were all performed in 14-day old adult females. What happens when they look at young female (like 2-day old). I assume that the nos>flp is working in larval and pupal stages and so the phenotype should be present in young females. Why did the authors choose this later age? For example, is the phenotype more robust in older females? or do you see more SGCs at later time points?

The authors did not supply any data to prove that the clones were larger in 14-day-old flies than in younger flies. Additionally, the age of "younger" flies was not specified. Therefore, the authors did not satisfactorily answer my concern.

We appreciate this critical comment. Figure 1J includes the SGC phenotype data from 1-, 7-, and 14-day-old flies. Both 1- and 7-day-old flies are younger flies in our analyses. The evidence that germline clones were larger in 14-day-old flies than in younger flies was provided in Figure 1-figure supplement 2 in the revised manuscript.

(6) Can the authors distinguish one copy of GFP versus 2 copies of GFP in germ cells of the ovary? This is not possible in the Drosophila testis. I ask because this could impact on the clonal analyses diagrammed in Fig. 4A and 4G and in 6A and B. Additionally, in most of the figures, the GFP is saturated so it is not possible to discern one vs two copies of GFP.

In the rebuttal, the authors stated that they cannot differential one vs two copies of GFP. They used other clone labeling methods in Fig. 4 and 6. I think that the authors should make a statement in the manuscript that they cannot distinguish one vs two copies of GFP for the record.

Thank you for this suggestion. Such statement has been added in the revised manuscript (Lines 177-178).

(7) More evidence is needed to support the claim of elevated Dpp levels in bam or bgcn mutant tumors. The current results with dpp-lacZ enhancer trap in Fig 5A,B are not convincing. First, why is the dpp-lacZ so much brighter in the mosaic analysis (A) than in the no-clone analysis (B); it is expected that the level of dpp-lacZ in cap cells should be invariant between ovaries and yet LacZ is very faint in Fig. 5B. I think that if the settings in A matched those in B, the apparent expression of dpp-lacZ in the tumor would be much lower and likely not statistically significantly. Second, they should use RNA in situ hybridization with a sensitive technique like hybridization chain reactions (HCR) - an approach that has worked well in numerous Drosophila tissues including the ovary.

The HCR FISH in Fig.5 of the revised manuscript needs an explanation for how the mRNA puncta were quantified. Currently, there is no information in the methods. What is meant but relative dpp levels. I think that the authors should report in and unbiased manner "number" of dpp or gbb puncta in TFs. For the germaria, I think that they should report the number of puncta of dpp or gbb divide by the total area in square pixels counted. Additionally, the background fluorescence is noticeably much higher in bamBG/delta86 germaria, which would (falsely) increase the relative intensity of dpp and gbb in bam mutants. Although, I commend the authors for performing HCR FISH, these data are still not convincing to me.

We appreciate these critical comments. Due to variable puncta sizes and frequent clustering in TF and cap cells (see Figure 5A, C), direct quantification of puncta number was unreliable. Therefore, we quantified mean fluorescence intensity instead, as described in the revised figure legend of Figure 5 (Lines 603-604). In Author response image 1 1A, B (modified from Figure 5A, C) , magenta ovals indicate empty background fluorescence areas, which appear similar between w1118 (wild-type control) and bam-/- germaria. In Author response image 1, the yellow oval outlines a neighboring germarium, not an empty area (see the DAPI channel).

In situ-HCR results of dpp and gbb in wild-type and bam mutant germaria.

Magenta ovals indicate empty areas displaying only background fluorescence. In panel (B), the yellow oval outlines a neighboring germarium, not an empty area (see the DAPI channel below).

(8) In Fig 6, the authors report results obtained with the bamBG allele. Do they obtain similar data with another bam allele (i.e., bamdelta86)?

The authors did not try any experiments with the bamdelta86 allele, despite this allele being molecularly defined, where the bamBG allele is not defined.

While we agree that repeating the experiments in Figure 6 with bamΔ86 would be helpful, our mosaic analysis strategy for two genes on different chromosome arms is technically complex (see genotypes in Source data 1). Switching from bamBG to bamΔ86 would necessitate extensive and time-consuming genetic recombination. Given that both alleles induce the SGC phenotype indistinguishably (Figure 1J), we believe that repeating these experiments with bamΔ86 would not alter our key conclusion. We appreciate your understanding regarding this technical complexity.

Reviewer #2 (Public review):

In the current version, Zhang et al. have made substantial improvements to the manuscript. It is now easier to read, and the data are more solid compared with the previous version, supporting their conclusion that tumor GSCs secrete stemness factors (BMPs and Dpp) to suppress the differentiation of neighboring wild-type GSCs. This study should benefit a broad readership across developmental biology, germ cell biology, stem cell biology, and cancer biology.

Thank you for your valuable comments, and we greatly appreciate them.

However, the following suggestions may further improve the clarity and rigor of the research content:

(1) Clarification of sample size (n).

Each germarium can contain highly variable numbers of SGCs, sometimes reaching 50-100. When reporting "n" values, the authors are encouraged to also indicate the number of germaria analyzed. For example, in lines 126-128:

"Notably, 74% of SGCs (n = 132) were GFP-negative, while the remaining 26% were GFP-positive (Figure 2B, C). This suggests that SGCs can be categorized into two distinct groups: those resembling GSCs (GSC-like) and those resembling cystoblasts (cystoblast-like)." Please clarify how many germaria were examined to obtain n = 132.

We appreciate this comment. In 14-day-old fly ovaries, each germarium that met our criterion for quantifying the SGC phenotype contains approximately 1.5 SGCs (see Figure 1K). For the specific analysis of the “132” SGCs presented in Figure 2C, we did not record the number of germaria from which they originated.

In addition, it is unclear whether the authors intend to suggest that the GFP-negative SGCs are GSC-like or cystoblast-like; this point should be clarified.

Thank you for this suggestion. We intend to suggest that the bamP-GFP-negative SGCs are GSC-like, which information has been added in the revised manuscript (Line 129).

(2) Improvement of Fig. 6 in situ hybridization images.

The in situ hybridization images in Fig. 6 are not fully convincing. The control images, in particular, would benefit from higher resolution and enlarged views of the germarium region.

Thank you for this valuable suggestion. The enlarged views of both the control and bam-/- germarium regions were included in Figure 5A, C in the revised manuscript.

In panel C, abundant signals are also present outside the germarium, which may complicate interpretation and should be clarified or controlled for.

In the right panel of Figure 5C, the abundant signals noted by the reviewer originate from neighboring germaria (see the DAPI channel), not from empty areas, which would be expected to show only background fluorescence. For more details, please refer to our response to Question (7) raised by Reviewer #1.

Alternatively, the authors could strengthen the in situ analysis by using bam mutants or bam dpp / bam gbb double mutants as controls to better define signal specificity.

We appreciate this comment. Homozygous dpp or gbb mutants are lethal, precluding the generation of dpp bam or gbb bam double-mutant flies. Additionally, the GFP signal was drastically reduced during our HCR processing, preventing mosaic clone analysis.

Reviewer #3 (Public review):

Zhang et al. investigated how germline tumors influence the development of neighboring wild-type (WT) germline stem cells (GSC) in the Drosophila ovary. They report that germline tumors generated by differentiation-arrested mutations (bam and bgcn) inhibit the differentiation of neighboring WT GSCs by arresting them in an undifferentiated state, resulting from reduced expression of the differentiation-promoting factor Bam. They find that these tumor cells produce low levels of the niche-associated signaling molecules Dpp and Gbb, which suppress bam expression and consequently inhibit the differentiation of neighboring WT GSCs non-cell-autonomously. Based on these findings, the authors propose that germline tumors mimic the niche to suppress the differentiation of the neighboring wild-type germline stem cells.

Strengths:

The study uses a well-established in vivo model to address an important biological question concerning the interaction between germline tumor cells and wild-type (WT) germline stem cells in the Drosophila ovary. If the findings are substantiated, this study could provide valuable insights that are applicable to other stem cell systems.

Thank you for your valuable comments, and we greatly appreciate them.

Weaknesses:

The authors have addressed some of my concerns in the revised submission. However, the data presented do not allow the authors to distinguish whether the failed differentiation of WT stem cells/germline cells results from "arrested differentiation due to the loss of the differentiation niche" or from "direct inhibition by tumor-derived expression of niche-associated molecules Dpp and Gbb".

Blocking Dpp or Gbb secretion specifically from germline tumor cells promoted differentiation of neighboring wild-type germ cells (Figure 6). This indicates that BMP ligands secreted by germline tumors are required to inhibit this differentiation. However, we cannot rule out the possibility that disruption of the differentiation niche also contributes to the SGC phenotype, a point highlighted in the manuscript (Line 204).

The critical supporting data, HCR in situ results, are not sufficiently convincing.

Below, we provide a point-by-point reply addressing each of your specific recommendations.

Recommendations for the authors:

Reviewer #3 (Recommendations for the authors):

It's a surprising that the authors failed to induce germline tumors at the adult stage, as this has been reported by many labs and would allow for time course analysis of SGC phenotype. As a result, the data in this manuscript address only events occurring after the germline tumor formation (with clonal induction at larval stage) and and focus on the already presene "arrested wild-type germ cells", without providing insight into the process of by which these arrested germ cells are formed.

In our hands, inducing germline clones by the hs-FLP method at the adult stage was efficient in males but not in females, despite subjecting adult flies to intensive heat-shock at 37°C.

The HCR in situ data exhibit a high background.

Regarding the background issue, please see our response to Reviewer #1’s Question (7).

First, the signal appears stronger in TF cells than in cap cells.

As demonstrated by Li et al. (Li et al., 2016), dpp-lacZ (P4-lacZ) signals are also stronger in TF cells than in cap cells (see their Figure 4D').

Second, both dpp and gbb are detected broadly in somatic cells including escort cells. These observations are inconsistent with published data.

As shown in Figure 5A and C, dpp and gbb were detected broadly in somatic cells of bam-/- germaria, but not in those of w1118 (wild-type) controls. To our knowledge, no previous study has reported the expression pattern of these ligands in a bam mutant background.

To demonstrate the tumor-derived dpp and gbb, the HCR in situ analysis could be performed in the germarium with mosaic clones. If these niche-associated molecules are indeed expressed in tumor cells, the authors should observe a mosaic expression pattern of these molecules, with signal "ON" in tumor cells and "OFF" in neighbouring arrested germ cells.

This is a great idea and was indeed our original approach. However, GFP signal was drastically reduced during our HCR processing, ultimately precluding mosaic clone analysis.

References

Chau, J., Kulnane, L.S., and Salz, H.K. (2009). Sex-lethal facilitates the transition from germline stem cell to committed daughter cell in the Drosophila ovary. Genetics 182, 121-132.

Li, X., Yang, F., Chen, H., Deng, B., Li, X., and Xi, R. (2016). Control of germline stem cell differentiation by Polycomb and Trithorax group genes in the niche microenvironment. Development 143, 3449-3458.

Li, Y., Minor, N.T., Park, J.K., McKearin, D.M., and Maines, J.Z. (2009). Bam and Bgcn antagonize Nanos-dependent germ-line stem cell maintenance. Proc Natl Acad Sci U S A 106, 9304-9309.

Changes from the earlier version to the revised version

Abstract

Although it is well-established that stem cells maintain tissue homeostasis while tumors disrupt it, the mechanisms by which tumors influence the development of nearby stem cells remain poorly understood. Using Drosophila ovaries as a model system, here we discovered that bam or bgcn mutant germline tumors inhibit the differentiation of neighboring wild-type germline stem cells (GSCs). Mechanistically, these tumor cells mimic the stem cell niche by secreting the BMP ligands Dpp and Gbb, but at reduced levels, resulting in moderate BMP signaling activation in adjacent GSCs. Such BMP signaling activation is sufficient to repress bam transcription, thereby blocking GSC differentiation. To our knowledge, this is the first example that tumors can functionally mimic a stem cell niche to inhibit the differentiation of neighboring wild-type stem cells. Similar regulatory paradigms may operate in mammalian tissues, including humans, during tumorigenesis.

Introduction

The homeostasis of many tissues in our bodies is maintained by adult stem cells, but this balance can be disrupted by tumor cells. What occurs when tumorigenesis intersects with stem cell development? To address this question, a mosaic-analysis model system is essential, where wild-type stem cells develop alongside tumor cells. Drosophila offers an exceptional model for such studies, as it allows for the efficient generation of mosaic clones through various established methods (Germani et al., 2018; Pastor-Pareja and Xu, 2013).

In Drosophila ovaries, germline stem cells (GSCs) play a crucial role in sustaining normal oogenesis and maintaining fertility (Fuller and Spradling, 2007; Lin, 1997). These GSCs reside in a specialized microenvironment known as the stem cell niche (hereafter referred to as niche) (Xie and Spradling, 2000). Typically, a GSC undergoes asymmetric division, generating two distinct daughter cells: one remains in the niche to self-renew as a GSC, while the other, called a cystoblast, exits the niche and initiates differentiation. During the differentiation process, each cystoblast performs exactly four rounds of mitotic division with incomplete cytokinesis to produce16 interconnected cystocytes, forming a germline cyst. In each germline cyst, only one germ cell is destined to become the oocyte, while the remaining 15 differentiate into nurse cells that support the development of the oocyte (Figure 1A) (Fuller and Spradling, 2007; Lin, 1997). The principal niche signals are bone morphogenetic protein (BMP) ligands, including Decapentaplegic (Dpp) and Glass bottom boat (Gbb), which are secreted by terminal filament (TF) and cap cells (Chen and McKearin, 2003a; Li et al., 2016; Song et al., 2004; Xie and Spradling, 1998, 2000). These ligands activate BMP signaling in GSCs, leading to the transcriptional repression of bag of marbles (bam), a key gene that promotes differentiation. In contrast, BMP signaling is inactive in cystoblasts, allowing Bam to be expressed and drive their differentiation (Chen and McKearin, 2003a; Song et al., 2004). Bam carries out this function in collaboration with its partner, Benign gonial cell neoplasm (Bgcn) (Li et al., 2009; Ohlstein et al., 2000).

bam or bgcn mutant germline tumors inhibit the differentiation of neighboring wild-type GSCs.

Earlier version: (A) Schematic cartoon for early oogenesis. The red dots and branches indicate spectrosomes and fusomes, respectively. TF cell: terminal filament cell; GSC: germline stem cell. (B) Mosaic analysis strategy. The FLP recombinase triggers mitotic recombination by targeting FRT sequences. The nos>FLP method restricts FLP expression to the germline, while the hs-FLP method enables heatshock-inducible FLP expression. (C-F) Representative samples. The asterisks mark cap cells, and the arrows indicate SGCs that have exited the niche and are surrounded by bam or bgcn mutant germline tumors. Vasa, a germ cell marker, should label all germ cells. However, due to poor tumor permeability, staining often fails to detect tumorous germ cells in the central region (see Vasa panels in D-F). (G-I) Representative samples (z-stack projections). In (G), the arrowheads and arrow respectively mark two GSCs and one cystoblast, all containing dot-like spectrosomes, while the dotted lines delineate cystocytes with branched fusomes. In (H) and (I), the arrows denote SGCs that also contain dot-like spectrosomes, akin to GSCs and the adjacent GSC-like tumor cells. (J and K) Quantification data. For each experiment, three independent replicates were performed, and data represent mean ± SEM. In (J), over 100 SGCs and germline cysts were quantified per replicate, and statistical significance was determined by one-way ANOVA. n.s. (P > 0.05). In (K), over 100 germaria were quantified per replicate. The online version of this article includes the following source data and figure supplement for figure 1: Source data 1. All genotypes. Source data 2. Raw quantification data. Figure supplement 1. SGCs appear in egg chambers.

Revised version: (A) Schematic cartoon for early oogenesis. The red dots and branches indicate spectrosomes and fusomes, respectively. TF cell: terminal filament cell; GSC: germline stem cell. (B) Mosaic analysis strategy. The FLP recombinase triggers mitotic recombination by targeting FRT sequences. The nos>FLP method restricts FLP expression to the germline, while the hs-FLP method enables heatshock-inducible FLP expression. (C-F) Representative samples. The asterisks mark cap cells, and the arrows indicate SGCs that have exited the niche and are surrounded by bam or bgcn mutant germline tumors. Vasa, a germ cell marker, should label all germ cells. However, due to poor tumor permeability, staining often fails to detect tumorous germ cells in the central region (see Vasa panels in D-F). (G-I) Representative samples (z-stack projections). In (G), the arrowheads and arrow respectively mark two GSCs and one cystoblast, all containing dot-like spectrosomes, while the dotted lines delineate cystocytes with branched fusomes. In (H) and (I), the arrows denote SGCs that also contain dot-like spectrosomes, akin to GSCs and the adjacent GSC-like tumor cells. (J and K) Quantification data. bamBG is strong loss-of-function allele of bam (Chen and McKearin, 2005). For each experiment, three independent replicates were performed, and data represent mean ± SEM. In (J), over 100 SGCs and germline cysts were quantified per replicate, and statistical significance was determined by one-way ANOVA. n.s. (P > 0.05). In (K), over 100 germaria were quantified per replicate.

Added in the revised version: SGCs appear in egg chambers.

Added in the revised version: (A, C) Representative samples. The arrows indicate SGCs enclosed within egg chambers. (B) Schematic of the experimental strategy for (C). See also Source data 1.

Added in the revised version: bam mutant germline clones enlarge as flies age.

Added in the revised version: (A) Representative samples. The dotted lines outline germline clones. All images are of the same magnification. (B) Quantification data. Over 30 germaria were quantified at each time point. Data represent mean ± SEM, and statistical significance was determined by one-way ANOVA. See also Source data 1 and 2.

Added in the revised version: Comparison of SGC phenotypes induced by the nos>FLP/FRT and hs-FLP/FRT systems.

Added in the revised version: (A) Representative samples. The arrows denote SGCs. All images are of the same magnification. (B) Quantification data. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed, and over 100 SGCs and germline cysts were quantified per replicate. Data represent mean ± SEM. See also Source data 1 and 2.

GSCs mutant for bam or bgcn fail to differentiate and instead hyper-proliferate, forming a well-established Drosophila germline tumor model (Lavoie et al., 1999; McKearin and Ohlstein, 1995; McKearin and Spradling, 1990; Niki and Mahowald, 2003). Notably, these germline tumor cells competitively displace wild-type GSCs from the niche (Jin et al., 2008). The resulting displacement creates a microenvironment where wild-type GSCs are surrounded by tumor cells, providing an excellent model system to study stem cell behavior in tumor neighborhoods.

Here, we demonstrate that bam or bgcn mutant germline tumors inhibit the differentiation of neighboring wild-type GSCs by functionally mimicking the stem cell niche. This mechanism may be conserved in mammals, including humans, during tumorigenesis, where malignant cells could similarly disrupt normal stem cell development.

Results

Germline tumors inhibit the differentiation of neighboring wild-type GSCs

To generate bam or bgcn mutant germline clones, we employed either nos>FLP/FRT or hs-FLP/FRT system that we previously established (Zhang et al., 2023; Zhao et al., 2018). These two systems induce the expression of FLP recombinase either germline-specifically (nos-GAL4-VP16/UASz-FLP) or via heat shock (hs-FLP). The expressed FLP recombinase targets the FRT sites to mediate mitotic recombination on homologous chromosome arms, generating adjacent GFP-negative (bam or bgcn mutant) and GFP-positive (wild-type) germ cell populations (Figure 1B). Remarkably, we observed that many wild-type germ cells located outside the niche retained a GSC-like single-germ-cell (SGC) morphology (Figure 1C, D), even when encapsulated within egg chambers (Figure 1—figure supplement 1). Under normal conditions, GSCs that exit the niche differentiate into interconnected germline cysts, where germ cells are linked rather than remaining as individual, isolated cells (Fuller and Spradling, 2007; Xie and Spradling, 2000). To rule out the possibility that the SGC phenotype is an artifact caused by GFP expression, we repeated the experiments using RFP and arm-lacZ as alternative mosaic-analysis markers. Consistent results were observed (Figure 1E, F), confirming that the phenotype is not attributable to GFP.

To further confirm that these SGCs exhibit GSC-like characteristics, we conducted anti-α-Spectrin immunofluorescent staining, a method that labels a germline-specific organelle known as the spectrosome in GSCs and cystoblasts, and the fusome in cystocytes. GSCs perform complete cell division, whereas cystocytes undergo incomplete cytokinesis, remaining interconnected through fusomes and ring canals. Consequently, spectrosomes appear as dot-like structures, while fusomes exhibit branched morphologies (Figure 1G) (Lin et al., 1994). To accurately capture the three-dimensional (3D) architecture of spectrosomes and fusomes, we acquired z-stack images using confocal microscopy. Strikingly, these SGCs displayed dot-like spectrosomes, closely resembling those observed in wild-type GSCs and bam or bgcn mutant GSC-like tumor cells (Figure 1H, I). We also considered the possibility that SGCs might arise through the dedifferentiation of the cystocytes in germline cysts surrounded by germline tumors. If this were the case, such cystocytes would initially undergo complete cell division, leaving behind midbodies as markers of the late cytokinesis stage. When visualized by anti-α-Spectrin immunofluorescence, midbody appears as a central sphere that is slightly connected to two larger flanking structures, resembling a variant of nunchucks (Mathieu et al., 2022). Notably, in our analyses of over 50 germline cysts surrounded by bam mutant germline tumors, none contained midbodies, suggesting that dedifferentiation is unlikely to be the primary mechanism responsible for the SGC phenotype. Together, these findings indicate that bam or bgcn mutant germline tumors inhibit the differentiation of neighboring wild-type GSCs.

Earlier version: To quantify the SGC phenotype, which requires the presence of both germline tumors and out-of-niche wild-type germ cells, we analyzed germaria containing both. In 14-day-old fly ovaries, 70% of germaria (432/618) met this criterion. We calculated the percentage of SGCs relative to the total number of SGCs and germline cysts, considering the out-of-niche germ cells that are either fully enclosed by germline tumors (e.g., the right SGC in Figure 1I and the marked germline cyst in Figure 2G) or in contact with wild-type germ cells or somatic cells on only one side (e.g., the left SGC in Figure 1I and the outlined germline cyst in Figure 4C). Notably, the SGC phenotype was consistent across the 14-day period analyzed (Figure 1J). For subsequent quantification, we used 14-day-old flies, taking advantage of the robust germline tumor growth at this age to ensure efficient analysis. In qualifying germaria, the average number of SGCs was approximately 1.5 (Figure 1K).

Revised version: To quantify the SGC phenotype, which requires the presence of both germline tumors and out-of-niche wild-type germ cells, we analyzed germaria containing both. In 14-day-old fly ovaries, 70% of germaria (432/618) met this criterion. We calculated the percentage of SGCs relative to the total number of SGCs and germline cysts, considering the out-of-niche germ cells that are either fully enclosed by germline tumors (e.g., the right SGC in Figure 1I and the marked germline cyst in Figure 2G) or in contact with wild-type germ cells or somatic cells on only one side (e.g., the left SGC in Figure 1I and the outlined germline cyst in Figure 4C). Notably, the SGC phenotype was consistent across the 14-day period analyzed (Figure 1J). For either 1-, 7-, or 14-day time point, we measured the sizes of bam mutant germline clones in over 30 germaria containing these clones. To estimate 3D clone size, we counted cell numbers within the maximal 2D cross-sectional area of each clone. Clones were larger in 14-day-old flies than in either 1- or 7-day-old flies (Figure 1—figure supplement 2). Therefore, we selected the 14-day time point for all subsequent analyses to maximize experimental efficiency. In qualifying germaria, the average number of SGCs was approximately 1.5 (Figure 1K). Furthermore, the SGC phenotypes induced by the nos>FLP/FRT and hs-FLP/FRT systems were indistinguishable (Figure 1—figure supplement 3). Given its simplicity and germline specificity, we primarily used the nos>FLP/FRT system in the following studies.

The inhibition of SGC differentiation depends on the lack of Bam expression.

Earlier version: (A) Representative sample. The arrowhead marks a BamC-positive 4-cystocyte germline cyst, while the arrows indicate BamC-negative SGCs. (B) Representative sample. The asterisk denotes cap cells, and the dotted circles outline bamP-GFP-negative GSCs. The solid circle marks a bamP-GFP-positive cystoblast. The arrow and arrowhead point to bamP-GFP-negative and -positive SGCs, respectively. (C) Quantification data. 14-day-old flies were used for the analyses. CBs: cystoblasts. (D) Schematic of the experimental strategy for (E-H). In “with hs-bam” flies (E and G), wild-type germ cells (both bam+/+ and bam+/-) carry the hs-bam transgene, while control “without hs-bam” flies (F) lack this element in their wild-type germ cells. (E-G) Representative samples. The arrows mark SGCs with dot-like spectrosomes, while the arrowhead indicates a 4-cystocyte germline cyst containing branched fusomes. (H) Quantification data. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM, and statistical significance was determined by t test. n.s. (P > 0.05). The online version of this article includes the following source data for figure 2: Source data 1. All genotypes. Source data 2. Raw quantification data.

Revised version: (A) Representative sample. The arrowhead marks a BamC-positive 4-cystocyte germline cyst, while the arrows indicate BamC-negative SGCs. (B) Representative sample. The asterisk denotes cap cells, and the dotted circles outline bamP-GFP-negative GSCs. The solid circle marks a bamP-GFP-positive cystoblast. The arrow and arrowhead point to bamP-GFP-negative and -positive SGCs, respectively. (C) Quantification data. 14-day-old flies were used for the analyses. CBs: cystoblasts. (D) Schematic of the experimental strategy for (E-H). In “with hs-bam” flies (E and G), wild-type germ cells (both bam+/+ and bam+/-) carry the hs-bam transgene, while control “without hs-bam” flies (F) lack this element in their wild-type germ cells. (E-G) Representative samples. The arrows mark SGCs with dot-like spectrosomes, while the arrowhead indicates a 4-cystocyte germline cyst containing branched fusomes. (H) Quantification data. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM, and statistical significance was determined by t test. n.s. (P > 0.05).

The inhibition of differentiation in SGCs relies on the lack of Bam expression

Earlier version: Given that Bam is the key factor promoting GSC differentiation (McKearin and Ohlstein, 1995; Ohlstein and McKearin, 1997), we were very curious about the expression of Bam in SGCs. At first, we assessed Bam protein levels using immunofluorescent staining with an anti-BamC antibody (McKearin and Ohlstein, 1995). Strikingly, none of the SGCs examined (n > 100) were BamC-positive (Figure 2A). Then, we analyzed bam transcription levels using a bamP-GFP reporter (Chen and McKearin, 2003b). 100% of GSCs within the niche (n = 153) were GFP-negative, while 98% of cystoblasts (n = 106) were GFP-positive (Figure 2B, C), confirming that bam transcription is associated with the initiation of GSC differentiation (McKearin and Ohlstein, 1995). Notably, 74% of SGCs (n = 132) were GFP-negative, while the remaining 26% were GFP-positive (Figure 2B, C). This suggests that SGCs can be categorized into two distinct groups: those resembling GSCs (GSC-like) and those resembling cystoblasts (cystoblast-like). The cystoblast-like SGCs may have already initiated their differentiation program toward becoming cystocytes. Since bam transcription initiates in cystoblasts (McKearin and Spradling, 1990) but Bam proteins accumulate predominantly in cystocytes (McKearin and Ohlstein, 1995), the Bam protein levels in these cystoblast-like SGCs are likely below the detection threshold at this early stage.

Revised version: Given that Bam is the key factor promoting GSC differentiation (McKearin and Ohlstein, 1995; Ohlstein and McKearin, 1997), we were very curious about the expression of Bam in SGCs. At first, we assessed Bam protein levels using immunofluorescent staining with an anti-BamC antibody (McKearin and Ohlstein, 1995). Strikingly, none of the SGCs examined (n > 100) were BamC-positive (Figure 2A). Then, we analyzed bam transcription levels using a bamP-GFP reporter (Chen and McKearin, 2003b). 100% of GSCs within the niche (n = 153) were GFP-negative, while 98% of cystoblasts (n = 106) were GFP-positive (Figure 2B, C), confirming that bam transcription is associated with the initiation of GSC differentiation (McKearin and Ohlstein, 1995). Notably, 74% of SGCs (n = 132) were GFP-negative (GSC-like), while the remaining 26% were GFP-positive (cystoblast-like) (Figure 2B, C). The cystoblast-like SGCs may have already initiated their differentiation program toward becoming cystocytes. Since bam transcription initiates in cystoblasts (McKearin and Spradling, 1990) but Bam proteins accumulate predominantly in cystocytes (McKearin and Ohlstein, 1995), the Bam protein levels in these cystoblast-like SGCs are likely below the detection threshold at this early stage.

Next, we asked whether ectopic expression of Bam can drive SGCs to differentiate. To address this, we established two experimental scenarios: one with the hs-bam element and one without as the control. In the hs-bam scenario (with hs-bam), GFP-positive germ cells are wild-type (carrying hs-bam), while GFP-negative cells are bam mutant (lacking hs-bam). In the control scenario (without hs-bam), GFP-positive cells are wild-type, and GFP-negative cells are bam mutant (Figure 2D, see genotypes in Source data 1). To induce ectopic Bam expression, 12-day-old female flies were subjected to heatshock treatment, which involved heating at 37°C for two hours, twice daily with a 6-hour interval, and over two consecutive days. In the absence of heatshock treatment, the percentage of SGCs in ovaries of both genotypes showed no significant difference at either 12 or 14 days (Figure 2E-H), indicating that the hs-bam element alone, without heatshock, does not affect the phenotype. However, following heatshock treatment, the percentage of SGCs in ovaries with hs-bam was markedly reduced compared to those without hs-bam (Figure 2E-H), suggesting that ectopic Bam expression can drive SGCs to differentiate. Collectively, these results support that the differentiation defects of SGCs are due to the lack of Bam expression.

SGCs retain moderate BMP signaling activation

Within the niche, BMP signaling functions to repress bam transcription to inhibit GSC differentiation (Chen and McKearin, 2003a; Song et al., 2004). To investigate BMP signaling activation in SGCs, we employed immunofluorescent staining for pMad, a well-characterized marker of BMP signaling activity (Kai and Spradling, 2003). Surprisingly, we observed undetectable pMad levels in all SGCs examined (n > 100) (Figure 3A, B). To investigate this further, we examined the activity of Dad-lacZ, a highly sensitive BMP signaling reporter known to be activated not only in GSCs but also in cystoblasts (Kai and Spradling, 2003; Song et al., 2004). Notably, 73% of SGCs were lacZ-positive (n = 107), a proportion lower than that of GSCs within the niche, which showed 100% lacZ positivity (n = 122) (Figure 3C, D). Furthermore, when comparing Dad-lacZ expression levels exclusively in lacZ-positive cells, we found that SGCs exhibited significantly lower expression levels than GSCs within the niche (Figure 3C, E). These findings indicate that BMP signaling is activated in SGCs but at lower levels than those in GSCs within the niche.

SGCs maintain lower BMP signaling levels than GSCs within the niche.

Earlier version: (A and B) Representative samples. The asterisks mark cap cells, arrowheads indicate pMad-positive GSCs, and arrows point to pMad-negative SGCs. (C) Representative samples. The asterisks denote cap cells, arrowheads mark Dad-lacZ-positive GSCs, and arrows highlight Dad-lacZ-positive SGCs. The dotted cycles outline one Dad-lacZ-negative SGC. (D and E) Quantification data. 14-day-old flies were used for the analyses. In (E), data represent mean ± SEM, and statistical significance was determined by t test. (F) Representative sample. The asterisk marks a cap cell, while the arrows indicate a BrdU+ GSC within the niche. (G) Representative sample. The arrow indicates a BrdU+ SGC surrounded by germline tumors. (H) Quantification data. 14-day-old flies were used for the analyses. Statistical significance was determined by chi-squared test. The online version of this article includes the following source data for figure 3: Source data 1. All genotypes. Source data 2. Raw quantification data.

Revised version: (A and B) Representative samples. The asterisks mark cap cells, arrowheads indicate pMad-positive GSCs, and arrows point to pMad-negative SGCs. (C) Representative samples. The asterisks denote cap cells, arrowheads mark Dad-lacZ-positive GSCs, and arrows highlight Dad-lacZ-positive SGCs. The dotted cycles outline one Dad-lacZ-negative SGC. (D and E) Quantification data. 14-day-old flies were used for the analyses. In (E), data represent mean ± SEM, and statistical significance was determined by t test. (F) Representative sample. The asterisk marks a cap cell, while the arrows indicate a BrdU+ GSC within the niche. (G) Representative sample. The arrow indicates a BrdU+ SGC surrounded by germline tumors. (H) Quantification data. 14-day-old flies were used for the analyses. Statistical significance was determined by chi-squared test.

Added in the revised version: BMP signaling inhibits SGC differentiation.

Added in the revised version: (A) Schematic of the experimental strategy for (B-F). Genotypes were unambiguously distinguished using a triple-color system (red, yellow, and green). (B-D) Representative samples. The dotted cycles mark an SGC, while the solid lines outline germline cysts containing differentiating cystocytes. (E and F) Quantification data. 14-day-old flies were used for the analyses. (G) Schematic of the experimental strategy for (H-J). (H and I) Representative samples. The dotted lines mark an SGC, while the solid lines outline a germline cyst containing differentiating cystocytes. (J) Quantification data. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM, and statistical significance was determined by t test.

Beyond maintaining Drosophila female GSCs in the niche, BMP signaling also promotes their division (Xie and Spradling, 1998). Since the activation levels of BMP signaling in SGCs were lower than those in GSCs within the niche, we hypothesized that SGCs would exhibit slower proliferation rates than GSCs. To test this hypothesis, we performed BrdU incorporation assays. The results revealed that only 4.5% of SGCs were BrdU-positive (n = 1034), a significantly lower proportion than the 7.8% observed in GSCs within the niche (n = 1337) (Figure 3F-H). These findings further corroborate the reduced activation of BMP signaling in SGCs relative to GSCs.

BMP signaling inhibits SGC differentiation

Earlier version: Then we investigated whether BMP signaling functions to inhibit SGC differentiation. The BMP type II receptor Punt and the co-Smad Medea (Med) are essential for maintaining GSC stemness within the niche (Xie and Spradling, 1998). To investigate whether they are also required to inhibit SGC differentiation, we established a genetic scenario, in which GFP+/+ RFP-/- germ cells are punt-/- or med-/-; GFP+/- RFP+/- germ cells are punt+/-bam+/- or med+/- bam+/- (similar to wild-type); and GFP-/- RFP+/+ germ cells are bam-/-. In control experiments (with no punt or med mutation), GFP+/+ RFP-/- germ cells are wild-type; GFP+/- RFP+/- germ cells are bam+/-(similar to wild-type); and GFP-/- RFP+/+ germ cells are bam-/- (Figure 4A, see genotypes in Source data 1). Strikingly, the proportion of punt-/- or med-/-SGCs relative to total SGCs was significantly lower than in controls (Figure 4B-E). Conversely, among punt-/- or med-/- germ cells meeting our established criteria for SGC phenotype quantification, germline cysts constituted a higher percentage compared to controls (Figure 4F). These results indicate that Punt and Med function to inhibit SGC differentiation.

Revised version: Then we investigated whether BMP signaling functions to inhibit SGC differentiation. The BMP type II receptor Punt and the co-Smad Medea (Med) are essential for maintaining GSC stemness within the niche (Xie and Spradling, 1998). Therefore, we sought to determine whether they are also required to inhibit SGC differentiation. However, because distinguishing one versus two copies of GFP proved difficult in our germline mosaic assays, we established a genetic scenario, in which GFP+/+ RFP-/- germ cells are punt-/- or med-/-; GFP+/- RFP+/- germ cells are punt+/- bam+/-or med+/- bam+/- (similar to wild-type); and GFP-/- RFP+/+ germ cells are bam-/-. In control experiments (with no punt or med mutation), GFP+/+ RFP-/- germ cells are wild-type; GFP+/- RFP+/- germ cells are bam+/- (similar to wild-type); and GFP-/- RFP+/+ germ cells are bam-/- (Figure 4A, see genotypes in Source data 1). Strikingly, the proportion of punt-/-or med-/- SGCs relative to total SGCs was significantly lower than in controls (Figure 4B-E). Conversely, among punt-/- or med-/- germ cells meeting our established criteria for SGC phenotype quantification, germline cysts constituted a higher percentage compared to controls (Figure 4F). These results indicate that Punt and Med function to inhibit SGC differentiation.

Removed in the revised version: BMP signaling inhibits SGC differentiation.

Removed in the revised version: (A) Schematic of the experimental strategy for (B-F). Genotypes were unambiguously distinguished using a triple-color system (red, yellow, and green). (B-D) Representative samples. The dotted cycles mark an SGC, while the solid lines outline germline cysts containing differentiating cystocytes. (E and F) Quantification data. 14-day-old flies were used for the analyses. (G) Schematic of the experimental strategy for (H-J). (H and I) Representative samples. The dotted lines mark an SGC, while the solid lines outline a germline cyst containing differentiating cystocytes. (J) Quantification data. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM, and statistical significance was determined by t test. The online version of this article includes the following source data for figure 4: Source data 1. All genotypes. Source data 2. Raw quantification data.

Mothers against dpp (Mad) is the primary transcription factor of BMP signaling, and it is also essential for GSC maintenance in Drosophila ovaries (Xie and Spradling, 1998). Unlike punt and med, which reside on the same chromosome arm (3R) as bam, mad is located on a separate chromosome arm (2L). To investigate whether Mad is required to inhibit SGC differentiation, we established a genetic scenario, in which GFP-/-germ cells are mad-/- and GFP+/+ germ cells are bam-/-. In control experiments (with no mad mutation), GFP-/- germ cells are wild-type and GFP+/+ germ cells are bam-/- (Figure 4G, see genotypes in Source data 1). Notably, mad mutation significantly decreased the SGC proportion relative to controls (Figure 4H-J). These results suggest that, like Punt and Med, Mad also plays a crucial role in suppressing SGC differentiation. Together, these findings demonstrate that BMP signaling contributes to inhibiting SGC differentiation, despite at reduced activation levels.

Germline tumors secret Dpp and Gbb

The formation of a differentiation niche by escort cells is required for GSC differentiation and is known to be disrupted by bam mutant germline tumors (Chen et al., 2022; Kirilly et al., 2011). Although this niche disruption could contribute to the SGC phenotype, an unaddressed question is the source of the BMP ligands (Dpp and Gbb) that maintain BMP signaling activation within SGCs. Guided by the Occam’s Razor principle, we focused on bam or bgcn mutant germline tumor cells, as the SGC phenotype is dependent on their presence. To assess the expression of dpp and gbb, we employed third-generation in situ hybridization chain reaction (HCR) (Choi et al., 2018). Successful detection was confirmed by prominent signal foci in TF and cap cells (Figure 5A, C). To enable quantitative comparison, all experiments and confocal imaging were performed under identical parameters. Signal intensity within bam mutant germline tumors and wild-type cystocytes was normalized to the signal in wild-type TF and cap cells. Strikingly, bam mutant germline tumor cells exhibited significantly elevated expression of both dpp and gbb compared to wild-type cystocytes (Figure 5A-D).

Germline tumors secret Dpp and Gbb.

Earlier version: (A and C) Representative samples. The asterisks denote TF or cap cells. The dotted lines highlight wild-type (WT) cystocytes, while the solid lines outline bam mutant germline tumor cells. (B and D) Quantification data for in situ-HCR assays. 14-day-old flies were used for the analyses, and over 10 samples were quantified for each experiment. The dpp or gbb expression levels in both WT cystocytes and bam mutant germline tumor cells were normalized to the average expression levels in WT TF and cap cells. (E and F) Quantification data for RT-qPCR assays. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed. Data represent mean ± SEM, and statistical significance in (B and D) was determined by one-way ANOVA and in (E and F) by t test. The online version of this article includes the following source data for figure 5: Source data 1. All genotypes. Source data 2. Raw quantification data.

Revised version: (A and C) Representative samples. The asterisks denote TF or cap cells. The dotted lines highlight wild-type (WT) cystocytes, while the solid lines outline bam mutant germline tumor cells. The magenta box areas are enlarged below. (B and D) Quantification data for in situ-HCR assays. 14-day-old flies were used for the analyses, and over 10 samples were quantified for each experiment. dpp and gbb expression levels were quantified by mean fluorescence intensity using ImageJ. For both wild-type cystocytes and bam mutant germline tumor cells, these levels were normalized to the average expression levels measured in wild-type TF and cap cells. (E and F) Quantification data for RT-qPCR assays. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed. Data represent mean ± SEM, and statistical significance in (B and D) was determined by one-way ANOVA and in (E and F) by t test.

To more sensitively assess Dpp and Gbb expression, we performed real-time quantitative PCR (RT-qPCR) analyses in bam or bgcn mutant ovaries, comparing samples with and without germline-specific knockdown of dpp or gbb. Detection of reduced transcript levels in knockdown conditions would confirm active expression of these genes in the respective genetic backgrounds. Consistent with the essential roles of these two genes in fly viability, ubiquitous knockdown using act-GAL4 with either dpp-RNAi or gbb-RNAi caused lethality, which results also validated the efficacy of these RNAi lines. Notably, germline-specific knockdown of dpp or gbb significantly reduced their transcript levels compared to yellow (y) or white (w) knockdown controls (Figure 5E, F). Collectively, these findings demonstrate that bam or bgcn mutant germline tumors secrete the BMP ligands, albeit at lower levels than TF and cap cells.

Dpp and Gbb secreted by germline tumors are required to inhibit SGC differentiation

Earlier version: Finally, we investigated whether Dpp and Gbb secreted by germline tumors are required to inhibit SGC differentiation. Using a previously established double-mutant mosaic-analysis strategy for two genes on different chromosomes (Zhang et al., 2024; Zhang et al., 2023), we generated dpp bam or gbb bam double-mutant germline clones using two dpp mutant alleles, dppd6, dppd12, and one gbb allele, gbb1 (Figure 6A, B, see genotypes in Source data 1). Heterozygotes in any of these alleles did not affect GSC maintenance, germ cell differentiation, and female fly fertility (Figure 6—figure supplement 1). However, both dpp bam and gbb bam double-mutant germline tumor cells exhibited reduced proliferation rates compared to bam single-mutant controls (Figure 6—figure supplement 2), indicating that autocrine BMP signaling promotes bam mutant tumor growth. As mentioned earlier, our evaluation focused on germ cells that have exited the niche and are surrounded by germline tumors to quantify the SGC phenotype. Thus, it raises the question of whether the extent of tumor encirclement (i.e., being surrounded by more or fewer tumor cells) influences the phenotype. To investigate this, we compared the SGC phenotype in bigger and smaller bam mutant germline tumors. Using confocal microscopy, we analyzed 70 germaria containing bam mutant germline clones that met our criteria for SGC phenotype quantification. For each bam mutant germline clone, we scanned and measured its maximal 2D cross-sectional area as a proxy for its 3D size. The 35 bigger and 35 smaller clones were categorized as ‘bigger’ and ‘smaller’ tumors, respectively. Strikingly, the SGC phenotype remained consistent between the two tumor groups (Figure 6—figure supplement 3), aligning with our earlier finding that this phenotype is stable over a 14-day period (Figure 1J), a timeframe sufficient for substantial germline tumor growth. These results suggest that direct contact between tumorous and wild-type germ cells, rather than tumor size, is the primary determinant of this phenotype.

Revised version: Finally, we investigated whether Dpp and Gbb secreted by germline tumors are required to inhibit SGC differentiation. Using a previously established double-mutant mosaic-analysis strategy for two genes on different chromosomes (Zhang et al., 2024; Zhang et al., 2023), we generated dpp bam or gbb bam double-mutant germline clones using two dpp mutant alleles, dppd6, dppd12, and one gbb allele, gbb1 (Figure 6A, B, see genotypes in Source data 1). Heterozygotes in any of these alleles did not affect GSC maintenance, germ cell differentiation, and female fly fertility (Figure 6—figure supplement 1). However, both dpp bam and gbb bam double-mutant germline tumor cells exhibited reduced proliferation rates compared to bam single-mutant controls (Figure 6—figure supplement 2), indicating that autocrine BMP signaling promotes bam mutant tumor growth. As mentioned earlier, our evaluation focused on germ cells that have exited the niche and are surrounded by germline tumors to quantify the SGC phenotype. Thus, it raises the question of whether the extent of tumor encirclement (i.e., being surrounded by more or fewer tumor cells) influences the phenotype. To investigate this, we compared the SGC phenotype in bigger and smaller bam mutant germline tumors. A total of 70 germaria containing bam mutant germline clones were analyzed using the same method described in Figure 1—figure supplement 2B. The 35 bigger and 35 smaller clones were categorized as “bigger” and “smaller” tumors, respectively. Strikingly, the SGC phenotype remained consistent between the two tumor groups (Figure 6—figure supplement 3), aligning with our earlier finding that this phenotype is stable over a 14-day period (Figure 1J), a timeframe sufficient for substantial germline tumor growth (Figure 1-figure supplement 2). These results suggest that direct contact between tumorous and wild-type germ cells, rather than tumor size, is the primary determinant of this phenotype.

Dpp and Gbb secreted by germline tumors are required to inhibit SGC differentiation.

Earlier version: (A) Schematic of the experimental strategy for (C-F). (B) Schematic of the experimental strategy for (G-I). (C-E, G, and H) Representative samples. The arrows mark SGCs containing dot-like spectrosomes, while the arrowheads denote germline cysts with differentiating cystocytes that possess branched fusomes. (F and I) Quantification data for the SGC phenotype. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM. Statistical significance in (F) was determined by one-way ANOVA and in (I) by t test. The online version of this article includes the following source data and figure supplement for figure 6: Source data 1. All genotypes. Source data 2. Raw quantification data. Figure supplement 1. Monoallelic deletion of dpp or gbb does not affect GSC maintenance, germ cell differentiation, and female fly fertility. Figure supplement 2. dpp bam or gbb bam double-mutant germline tumor cells divide more slowly than bam single-mutant ones. Figure supplement 3. The SGC phenotype is unchanged irrespective of the number of surrounding germline tumors.

Revised version: (A) Schematic of the experimental strategy for (C-F). (B) Schematic of the experimental strategy for (G-I). (C-E, G, and H) Representative samples. The arrows mark SGCs containing dot-like spectrosomes, while the arrowheads denote germline cysts with differentiating cystocytes that possess branched fusomes. (F and I) Quantification data for the SGC phenotype. 14-day-old flies were used for the analyses. For each experiment, three independent replicates were performed, with over 100 SGCs and germline cysts quantified per replicate. Data represent mean ± SEM. Statistical significance in (F) was determined by one-way ANOVA and in (I) by t test.

Added in the revised version: Monoallelic deletion of dpp or gbb does not affect GSC maintenance, germ cell differentiation, and female fly fertility.

Added in the revised version: (A-D) Representative samples. (E) Quantification data. 14-day-old flies were used for the analyses. GSCs are germ cells that are located within the niche and contain dot-like spectrosomes. Cystoblasts are germ cells that have exited the niche, remain in contact with GSCs, and maintain dot-like spectrosomes. Cystocytes in germline cysts are germ cells that are characterized by branched fusomes. (F) Fertility test. 3-day-old flies were used for the analyses. For each genotype, three independent replicates were performed. Data represent mean ± SEM, and statistical significance was determined by one-way ANOVA. n.s. (P > 0.05). See also Source data 1 and 2.

Added in the revised version: dpp bam or gbb bam double-mutant germline tumor cells divide more slowly than bam single-mutant ones.

Added in the revised version: (A, B, D, and E) Representative samples. The arrows indicate BrdU+ germline tumor cells mutant for bam, dpp bam, or gbb bam. (C and F) Quantification data. 14-day-old flies were used for the analyses. Statistical significance was determined by chi-squared test. See also Source data 1 and 2.

Added in the revised version: The SGC phenotype is unchanged irrespective of the number of surrounding germline tumors.

Added in the revised version: (A) Representative samples. The arrows denote SGCs. All images are of the same magnification. (B) Quantification data for tumor size. Data represent mean ± SEM. (C) Quantification data for the SGC phenotype. In (B and C), 14-day-old flies were used for the analyses. Statistical significance in (B) was determined by t test and in (C) by chi-squared test. n.s. (P > 0.05). See also Source data 1 and 2.

The results above demonstrate that comparing the severity of the SGC phenotype is feasible between germ cells surrounded by smaller dpp bam or gbb bam double-mutant germline tumors and those surrounded by larger bam single-mutant germline tumors. Remarkably, both dpp bam and gbb bam double-mutant germline tumors enclosed fewer SGCs but more germline cysts than their bam single-mutant counterparts (Figure 6C-I). Thus, we concluded that the BMP ligands from directly-contacting germline tumor cells mediate the dominant inhibition of SGC differentiation. This amazingly parallels the mechanism observed in normal stem cell niche, where only germ cells in direct contact with cap cells are maintained as GSCs (Chen and McKearin, 2003a; Song et al., 2004; Xie and Spradling, 2000).

Discussion

Our study reveals that bam or bgcn mutant germline tumors in Drosophila ovaries secrete lower levels of BMP ligands Dpp and Gbb than cap cells, resulting in moderate BMP signaling activation in adjacent wild-type GSCs (called SGCs in this study). Such BMP signaling activation is sufficient to repress bam transcription, thereby blocking SGC differentiation (see our working model in Figure 7). Strikingly, this mechanism closely recapitulates the normal niche signaling program mediated by TF and cap cells (Chen and McKearin, 2003a; Song et al., 2004; Xie and Spradling, 1998, 2000). To our knowledge, this represents the first evidence that tumor cells can functionally mimic a stem cell niche to arrest neighboring wild-type stem cells in an undifferentiated state.

A working model.

bam or bgcn mutant germline tumors secrete the BMP ligands Dpp and Gbb to activate BMP signaling in out-of-niche GSCs (called SGCs in this study) to inhibit their differentiation (left panel). In contrast, dpp bam and gbb bam double-mutant germline tumors exhibit a significant loss of this differentiation-inhibiting ability (right panel).

While bam or bgcn mutant germline tumors consist of GSC-like cells expected to resemble SGCs (Lavoie et al., 1999; McKearin and Ohlstein, 1995), we found key differences in BMP signaling. Out-of-niche bgcn mutant tumor cells showed significantly lower BMP activity than neighboring SGCs, as evidenced by reduced Dad-lacZ expression (Figure 3C). Consistent with this, most of out-of-niche bam mutant tumor cells expressed bamP-GFP, a reporter suppressed by BMP signaling (Chen and McKearin, 2003a; Song et al., 2004), whereas only 26% of SGCs were positive (Figure 2B, C). These findings suggest that SGCs are more responsive to BMP signals secreted by germline tumors than the tumors themselves. Future studies are needed to elucidate the underlying mechanisms.

One interesting finding is that bam or bgcn mutant germline tumors secrete lower levels of BMP ligands than TF and cap cells (Figure 5A-D). This aligns with earlier microarray data showing that purified Drosophila female GSCs express minimal Dpp and Gbb (Kai et al., 2005). However, our work reveals that such BMP levels in germline tumors are functionally critical to suppress SGC differentiation (Figure 6). Unlike normal GSCs, which receive unidirectional BMP ligands from cap cells (Chen and McKearin, 2003a; Li et al., 2016; Song et al., 2004; Xie and Spradling, 2000), SGCs are often fully surrounded by bam or bgcn mutant germline tumors. This spatial advantage likely enables tumors to inhibit SGC differentiation efficiently without matching the high BMP output of TF and cap cells. Moreover, since BMP signaling is known to both inhibit normal GSC differentiation and promote their proliferation (Xie and Spradling, 1998), it should similarly stimulate SGC expansion, which is detrimental for bam or bgcn mutant germline tumors. We propose that these tumor cells finely regulate BMP secretion to balance these opposing demands: maintaining differentiation blockade of SGCs while avoiding stimulation of their excessive proliferation.

A well-established principle in oncology is that tumor aggressiveness correlates with poor differentiation, with less-differentiated tumors exhibiting enhanced transformative capacity and metastatic potential (Jogi et al., 2012; Lytle et al., 2018). In Drosophila ovaries, bam or bgcn mutant germline tumors consist of GSC-like cells that may resemble these poorly differentiated human tumors (Lavoie et al., 1999; McKearin and Ohlstein, 1995). This similarity raises the possibility that stem cell-like human tumors may similarly inhibit the differentiation of adjacent wild-type stem cells. By blocking differentiation, such tumors could deplete terminally differentiated cell populations, potentially exacerbating patient mortality. This mechanism may contribute to the heightened lethality of poorly differentiated tumors. Further investigation is needed to test this hypothesis.

The differentiation of a single GSC into a 16-cell germline cyst, comprising 15 polyploid nurse cells and one developing oocyte, represents a substantial metabolic investment (Fuller and Spradling, 2007; Lin, 1997). We propose that bam or bgcn mutant germline tumors block this process to divert nutrients toward their own uncontrolled growth. This phenomenon could have broad implications, as many human tissues and organs (intestine, muscle, skin, blood system, male germline, etc.) similarly depend on adult stem cells for homeostasis (Blanpain and Fuchs, 2006; Gehart and Clevers, 2019; Sousa-Victor et al., 2022; Spradling et al., 2011; Wilkinson et al., 2020). Notably, these stem cell-dependent tissues and organs are frequent sites of tumorigenesis, raising the possibility that human cancers may similarly impair neighboring stem cell differentiation to optimize nutrient allocation for malignant growth. A key limitation of our study is that the evidence is derived solely from Drosophila germline. Future work should explore whether similar regulatory paradigms operate in mammalian tissues during tumorigenesis.

Materials and Methods

Fly husbandry

Flies were raised at 25°C on standard cornmeal/molasses/agar media.

Transgenic flies

hs-bam on chromosome 3R: The coding sequence of the bam gene, amplified from the cDNA clone, was cloned into the BglII-XbaI sites of the pCaSpeR-hs vector, while the attB sequence was inserted into the XhoI site. The resulting attB-pCaSpeR-hs-bam plasmid was then microinjected into the attP154 (Chromosome 3R, 97D2) fly strain to generate site-specific transgenic flies.

Heatshock method to induce germline clones

To ensure developmental synchrony and maintain low-density growth, eggs within 8 hours of laying were collected for heatshock treatment. The animals (late-Larva 3/early-Pupa stage) were subjected to twice-daily heatshocks at 37°C (2 hours per session, with a 6-hour interval between the two sessions) for 6 consecutive days.

Fertility test

For each genotype, three independent crosses were performed. Each cross vial contained two females and four w1118 (wild-type) males, all aged three days old. The crosses were transferred to fresh vials every two days, with five replicate vials quantified per genotype. After all adult flies eclosed, offspring production was assessed by counting the number of empty pupae on the vial walls.

BrdU labeling

Ovaries were dissected in Schneider’s insect medium (SIM) and incubated in freshly prepared BrdU solution (100 μg/mL in SIM) for five hours at 25°C. After washing with PBS for 30 min, samples were fixed in 4% paraformaldehyde (in PBS) for three hours, followed by another PBS wash for 30 min. Samples were then treated with RQ1 DNase reaction solution (Promega, Madison, WI, USA) for one hour, washed with PBST (0.3% Triton X-100 in PBS) for 30 min, and incubated overnight at 4°C with mouse anti-BrdU antibody. Following a PBST wash for one hour, ovaries were incubated with goat anti-mouse 546 and DAPI (0.1 μg/mL) in PBST for three hours, washed again in PBST for one hour, and mounted in autoclaved 70% glycerol.

Immuno-florescent staining, image collection, and data Processing

Earlier version: Ovaries were dissected in PBS, fixed in 4% paraformaldehyde (in PBS) for three hours, washed with PBST for 30 min, and then incubated overnight at 4°C with primary antibodies diluted in PBST as follows: mouse anti-BamC (1:5), mouse anti-β-Gal (1:200), mouse anti-α-Spectrin (1:100), rabbit anti-pMad (1:500), and rabbit anti-Vasa (1:2000). After washing with PBST for one hour, samples were incubated with Alexa Fluor-conjugated secondary antibodies (1:1000) and 0.1 μg/mL DAPI (in PBST) for three hours, followed by a final PBST wash for one hour. Ovaries were mounted in autoclaved 70% glycerol and imaged using a Zeiss LSM 710 confocal microscope (Carl Zeiss AG, Baden-Württemberg, Germany). Images were processed with ZEN 3.0 SR imaging software (Carl Zeiss) and Adobe Photoshop 2022. The quantification data were processed by GraphPad Prism, ImageJ, or Microsoft Excel.

Revised version: Ovaries were dissected in PBS, fixed in 4% paraformaldehyde (in PBS) for three hours, washed with PBST for 30 min, and then incubated overnight at 4°C with primary antibodies. The rabbit anti-pMad antibody was a gift from Ed Laufer, and the rabbit anti-Vasa antibody was a gift from Zhaohui Wang (Chen et al., 2014). After washing with PBST for one hour, samples were incubated with Alexa Fluor-conjugated secondary antibodies and 0.1 μg/mL DAPI (in PBST) for three hours, followed by a final PBST wash for one hour. Ovaries were mounted in autoclaved 70% glycerol and imaged using a Zeiss LSM 710 confocal microscope (Carl Zeiss AG, Baden-Württemberg, Germany). Images were processed with ZEN 3.0 SR imaging software (Carl Zeiss) and Adobe Photoshop 2022. The quantification data were processed by GraphPad Prism, ImageJ, or Microsoft Excel.

In situ hybridization chain reaction (HCR)

Ovaries dissected from 14-day-old female flies were processed according to the following protocol.

Fixation: Ovaries were fixed in 4% paraformaldehyde (in PBS) for 3 hours at room temperature (RT) or overnight at 4°C;

Earlier version: Hybridization: Following fixation, samples were washed three times for 5 min each in PBST, dehydrated in methanol for 5 min, rehydrated through a methanol:PBST gradient series (3:1, 1:1, and 1:3), followed by another three 5-min PBST washes, treated with Proteinase K (10Lμg/mL) for 5 min, washed again three times for 5 min each in PBST, pre-hybridized in preheated hybridization buffer [50% formamide, 5× SSC, 9 mM citric acid (pH 6.0), 0.1% Tween 20, 50 µg/mL heparin, 1× Denhardt’s solution, 10% dextran sulfate] for 30 min at 37°C, and then incubated with RNA-FISH probes (0.1LμM in hybridization buffer) overnight at 37°C.

Revised version: Hybridization: Following fixation, samples were washed three times for 5 min each in PBST, dehydrated in methanol for 5 min, rehydrated through a methanol:PBST gradient series (3:1, 1:1, and 1:3), followed by another three 5-min PBST washes, treated with Proteinase K (10Lμg/mL) for 5 min, washed again three times for 5 min each in PBST, pre-hybridized in preheated hybridization buffer [50% formamide, 5× SSC, 9 mM citric acid (pH 6.0), 0.1% Tween 20, 50 µg/mL heparin, 1× Denhardt’s solution, 10% dextran sulfate] for 30 min at 37°C, and then incubated with in situ-HCR probes (0.1LμM in hybridization buffer) overnight at 37°C. For the detection of dpp and gbb, a pool of 20 in situ-HCR probes targeting each mRNA was employed. The probe sequences were provided in the Key Resources Table.

Earlier version: Signal amplification: The next day, samples were washed four times for 15 min each at 37°C with preheated probe wash buffer (50% formamide, 5× SSC, 9 mM citric acid, 0.1% Tween 20, 50 µg/mL heparin), followed by three 10-min washes in 5× SSCT (5× SSC, 0.1% Tween 20) at RT. After pre-hybridization in amplification buffer (5× SSC, 0.1% Tween 20, 10% sodium sulfate) for 10 min at RT, an amplification reaction was performed using heat-denatured hairpin h1/h2 nucleic acids (30 nM for each in amplification buffer) overnight in the dark at RT.

Revised version: Signal amplification: The next day, samples were washed four times for 15 min each at 37°C with preheated probe wash buffer (50% formamide, 5× SSC, 9 mM citric acid, 0.1% Tween 20, 50 µg/mL heparin), followed by three 10-min washes in 5× SSCT (5× SSC, 0.1% Tween 20) at RT. After pre-hybridization in amplification buffer (5× SSC, 0.1% Tween 20, 10% sodium sulfate) for 10 min at RT, an amplification reaction was performed using heat-denatured hairpin nucleic acids (30 nM for each in amplification buffer) overnight in the dark at RT. The hairpin sequences were provided in the Key Resources Table.

Washing and mounting: Samples were washed three times for 10 min each in 5× SSCT, followed by three 10-min washes in PBST, and then mounted in autoclaved 70% glycerol for imaging.

Removed in the revised version: For the detection of dpp and gbb, a pool of 20 RNA-FISH probes targeting each mRNA was employed. The sequences of these probes were as follows:

Removed in the revised version: dpp RNA-FISH probes:

Removed in the revised version: GAGGAGGGCAGCAAACGGaaAGAGCATGGCCACGCTGTCCAGTTG

Removed in the revised version: GCACCACCGTACTTTGGTCGTTGAGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaTCGTAGCCCAGAGGCGCCACAATCC

Removed in the revised version: GGGCACTTCCCGTGGCAGTAATATGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaTCTCGGCTGCCGCTTGTTCCGGCCG

Removed in the revised version: GTCGTGGTTCTTGCGCCTCGTAGGCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGCTGCTTGTGCTGCCACCGCTCGTG

Removed in the revised version: CGTCGTCCGTGTAGGTGAACAGGAGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGGCCGGCTGGACATCGAGGCTCACC

Removed in the revised version: CTGCGGACTCGCCAGCCACCGGTCCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaCACCTGGTAGCGCGTCCGATTCGCC

Removed in the revised version: CCCGACGCGCGTGATGTCGTAGACAtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaTGCTCTTCACGTCGAAGTGCAGCCG

Removed in the revised version: CCGCCTTCAGCTTCTCGTCGGCGGGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaCCCATGATCTCGGCGTAGAGCTTCT

Removed in the revised version: GGGATGTTGACCGAGTCGAGCTCGTtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaAGCGCCTCCTTGCTGTAGGTGGACG

Removed in the revised version: GGGTCTGGCTTCAGCTTGTCCTTGAtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaCACGAAGATTGATTCAATCGACGAG

Removed in the revised version: GCGGTCGAGCACCAGCGTCGGCTCCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: gbb RNA-FISH probes:

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGGTGGTACAGAACGGGTAGTGCTCC

Removed in the revised version: TTTTCAGGTTCACATTCTCGTCGTTtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGCGTTCATGTGCGCATTGAGCGGGA

Removed in the revised version: AGGGTCTGGACGATCGCATGGTTCGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGTACAGGGTCTGCATCTGGCAGCTG

Removed in the revised version: ATGCCAGCCCAGATCCTTGAAGTCTtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaTGCTCCTGTGGTGGCTGCTGTGGGC

Removed in the revised version: GCTTGCGTGGATGGCTGGCGCTTCGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaATGTCGTCCAGCTTCACCTCGCGGT

Removed in the revised version: TCGTCCACCTTGCGGTGGATCAGTCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGTTGAGCTCCAACCAGCCCACGTAG

Removed in the revised version: CAGCCACTCGTGCAGGCCCTCGGTCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaACTCCCTGTTGGCGGTCAGCCACTT

Removed in the revised version: TGCCAATGGCGTATACCGTGATGGTtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaCGACGGCCGTGCTCGTGACGCAGTT

Removed in the revised version: GGCACGTTGGAGACGTCGAACCACAtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaGTTCTTCTGCTGCTCGCCCTCATCC

Removed in the revised version: GGCCCGCTTGTCCAGGTCGGTGATGtaGAAGAGTCTTCCTTTACG

Removed in the revised version: GAGGAGGGCAGCAAACGGaaTGATGCGGTGGTAGACGTCCAGCAG CCTGATCGCTGAGACCCTCCTCCGCtaGAAGAGTCTTCCTTTACG

Removed in the revised version: The hairpin h1/h2 sequences were as follows:

Removed in the revised version: B1H1-594:

Removed in the revised version: CGTAAAGGAAGACTCTTCCCGTTTGCTGCCCTCCTCGCATTCTTTCTTGAGGAGGG

Removed in the revised version: CAGCAAACGGGAAGAG

Removed in the revised version: B1H2-594:

Removed in the revised version: GAGGAGGGCAGCAAACGGGAAGAGTCTTCCTTTACGCTCTTCCCGTTTGCTGCCCT

Removed in the revised version: CCTCAAGAAAGAATGC

Real-time quantitative PCR (RT-qPCR)

Earlier version: Ovaries from 14-day-old flies were dissected, and total RNA was extracted using the RNeasy Micro Kit. Equal amounts of RNA were reverse-transcribed into cDNA using the HiFiScript cDNA Synthesis Kit. RT-qPCR was performed on a CFX Connect Real-Time PCR System (Bio-Rad) with ChamQ SYBR qPCR Master Mix. The PCR protocol consisted of an initial denaturation at 95°C for 30 min, followed by 40 cycles of 95°C for 10 sec and 60°C for 30 sec. Relative gene expression was calculated using the 2−ΔΔCT method (Livak and Schmittgen, 2001). The primers described previously were used (Huang et al., 2017):

Revised version: Ovaries from 14-day-old flies were dissected, and total RNA was extracted using the RNeasy Micro Kit. Equal amounts of RNA were reverse-transcribed into cDNA using the HiFiScript cDNA Synthesis Kit. RT-qPCR was performed on a CFX Connect Real-Time PCR System (Bio-Rad) with ChamQ SYBR qPCR Master Mix. The PCR protocol consisted of an initial denaturation at 95°C for 30 min, followed by 40 cycles of 95°C for 10 sec and 60°C for 30 sec. Relative gene expression was calculated using the 2−ΔΔCT method (Livak and Schmittgen, 2001). The primers used, which were previously described (Huang et al., 2017), were listed in the Key Resources Table.

Removed in the revised version: dpp primer-1: TACCACGCCATCCACTCAAC

Removed in the revised version: dpp primer-2: GCTCGTTACTCGATACGGCT

Removed in the revised version: gbb primer-1: CTGGATCATCGCACCAGAGG

Removed in the revised version: gbb primer-2: GTCTGGACGATCGCATGGTT

Removed in the revised version: rp49 (internal control) primer-1: CACCGGATTCAAGAAGTTCC

Removed in the revised version: rp49 (internal control) primer-2: GACAATCTCCTTGCGCTTCT